---
title       : Accessing REDCap from R
description : How to use the REDCap API





--- type:NormalExercise lang:r xp:100 skills:1 key:e7e6596826
## Accessing 



*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
