---
title       : Understanding Psychometric Instruments at CAR
description : CAR has a ton of instruments that we use to measure all sorts of things!  You'll want to be familiar with many of these.


--- type:NormalExercise lang:r xp:100 skills:1 key:0f35797452
## Access to CAR's form library

Hopefully, if you were assigned this training, your PI (or Joy) has added you to the following database.  See if you can access it:
https://redcap.research.chop.edu/redcap_v7.1.2/index.php?pid=12524

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
