---
title       : Understanding Psychometric Instruments at CAR
description : CAR has a ton of instruments that we use to measure all sorts of things!  You'll want to be familiar with many of these.


--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:03418b5e2c
## Access to CAR's form library

Hopefully, if you were assigned this training, your PI (or Joy) has added you to the following database.  See if you can access it:
https://redcap.research.chop.edu/redcap_v7.1.2/index.php?pid=12524.  You can only advance in this chapter once you have access.  If you don't, [email Joy](mailto:paytonk@email.chop.edu)!

Ready to move on?

*** =instructions
- Yep, I have access!
- Nope, still waiting!

*** =hint
No real hint.  Just ask Joy to add you!

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}
msg_bad <- "Oops, you have to get that access before you can continue!"
msg_good <- "Alright, let's keep going!"
test_mc(correct = 5, feedback_msgs = c(msg_good, msg_bad))
```
